# Lab 0: Introduction to MDS homework and practice with Jupyter Notebooks

In Lab 0, you will learn how to navigate a Jupyter Notebook (like this one). You should follow the instructions in the notebook to complete the assignment.

### Submission Instructions
rubric={mechanics:2}

- [Follow the submission instruction on Canvas](https://canvas.ubc.ca/courses/53662/assignments/635976)

## 1. Jupyter notebooks
This webpage is called a Jupyter notebook. A notebook is a place to write programs and view their results.

### 1.1. Text cells
In a notebook, each rectangle containing text or code is called a *cell*.

Text cells (like this one) can be edited by double-clicking on them. They're written in a simple format called [Markdown](http://daringfireball.net/projects/markdown/syntax) to add formatting and section headings.  You will learn more about Markdown as you go through MDS and you will use it often! 

After you edit a text cell, select the "run cell" button at the top that looks like ▶ to confirm any changes.

### Question 1.1.1.
rubric={correctness:1}

This paragraph is in its own text cell.  Try editing it so that **this** sentence is the last sentence in the paragraph, and then select the "run cell" ▶ button on the top.  This sentence, for example, should be deleted.  So should this one.

Now try running the code in the cells below, again using the "run cell" ▶ button.

```python
print("Hello, World!")
```

You can also run cells using the keyboard shortcuts `Shift+Enter`. Try using that keyboard shortcut to run the cell below: 

```python
print("\N{WAVING HAND SIGN}")
```

You can also print many emojis by pasting them directly.

```python
print("👋")
```

The fundamental building block of Python code is an expression. Cells can contain multiple lines with multiple expressions. When you run a cell, the lines of code are executed in the order in which they appear. Every `print` expression prints a line. Run the next cell and notice the order of the output.

```python
print("First this line is printed,")
print("then this one,")
```

### Question 1.1.2.
rubric={correctness:1}

Change the cell below by replacing the following code:

```
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```

with another `print` expression so that it prints out:

    First this line is printed,        
    then this one,
    and then the whole 🌏.

*Hint:* If you're stuck on how to print the Earth symbol, try looking at the print expressions above.

<!-- grade_id: cell-f7e5cc91909e08b4, points: 0, grade: true, solution: true -->
```python
print("First this line is printed,")
print("then this one,")
# YOUR CODE HERE
raise NotImplementedError()
```

### 1.2. Comments
You may have noticed this line in the cell above:

```
# your code here
```

That is called a comment. It doesn't make anything happen in Python; Python ignores anything on a line after a `#`. Instead, it's there to communicate something about the code to you, the human reader. Comments are extremely useful. Comments also exist in R and use the same `#` symbol.

![](https://imgs.xkcd.com/comics/future_self.png)

*source: https://xkcd.com/1421/*


### 1.3. Writing Jupyter notebooks
You can use Jupyter notebooks for your own projects or documents. They are among the world's most popular programming environments for data science. When you make your own notebook, you'll need to create your own cells for text and code.

To add a cell, select the + button in the menu bar.  You can change between code cell and text cell by selecting the cell so that it's highlighted, then selecting the drop-down box next to the restart (⟳) button in the menu bar, and choosing Code instead of Markdown.

### Question 1.3.1.
rubric={correctness:1}

Add a code cell below this one.  Write code in it that prints out:
   
    A whole new cell! 
    ♪🌏♪

(That musical note symbol is like the Earth symbol.  Its long-form name is `\N{EIGHTH NOTE}`.)

Run your cell to verify that it works.

<!-- grade_id: cell-08528d6dc121aca7, points: 0, grade: true, solution: true -->
YOUR ANSWER HERE

### 1.4. Errors
Python is a language, and like natural human languages, it has rules.  It differs from natural language in two important ways:
1. The rules are *simple*.  You can learn most of them in a few weeks and gain reasonable proficiency with the language in a semester.
2. The rules are *rigid*.  If you're proficient in a natural language, you can understand a non-proficient speaker, glossing over small mistakes.  A computer running Python code is not smart enough to do that.

Whenever you write code, you'll make mistakes.  When you run a code cell that has errors, Python will sometimes produce error messages to tell you what you did wrong.

Errors are okay; even experienced programmers make **many** errors.  When you make an error, you just have to find the source of the problem, fix it, and move on.

We have made an error in the next cell.  Run it and see what happens.

```python
print("This line is missing something."
```

You should see something like this (minus our annotations):

![](imgs/error.jpg)

The last line of the error output attempts to tell you what went wrong.  The *syntax* of a language is its structure, and this `SyntaxError` tells you that you have created an illegal structure.  "`EOF`" means "end of file," so the message is saying Python expected you to write something more (in this case, a right parenthesis) before finishing the cell.

There's a lot of terminology in programming languages. You'll learn as you go. If you are ever having trouble understanding an error message, search the discussion forum. If you don't find an answer, post a question about the error yourself.

Try to fix the code above so that you can run the cell and see the intended message instead of an error.

### 1.5. The Kernel
The kernel is a program that executes the code inside your notebook and outputs the results. In the top right of your window, you can see a circle that indicates the status of your kernel. If the circle is empty (⚪), the kernel is idle and ready to execute code. If the circle is filled in (⚫), the kernel is busy running some code. 

You may run into problems where your kernel is stuck for an excessive amount of time, your notebook is very slow and unresponsive, or your kernel loses its connection. If this happens, try the following steps:
1. At the top of your screen, select **Kernel**, then **Interrupt**.
2. If that doesn't help, select **Kernel**, then **Restart**. If you do this, you will have to run your code cells from the start of your notebook up until where you paused your work.
3. If that doesn't help, restart your server. But first, save your work by selecting **File** at the top left of your screen, then **Save**.

Jupyter can run many different kernels, meaning that you can run a Python kernel to use Python inside a Jupyter notebook AND you can also run a R kernel to use R inside a Jupyter notebook. When you are inside a Jupyter notebook the name of the kernel can be found beside the circle that tells you what your kernel is doing (⚪). This notebook is using the Python 3 kernel. 

### 1.6. Autograded questions

Many assignments in MDS will be distributed as notebooks like this one. At the top of each assignment, you'll see a cell like the one below that imports the code needed to run the autograder tests. Run it to import the autograder tests.



```python
# Don't change this cell, just run it
# Import code needed for autograder tests
from hashlib import sha1
```

Autograded questions start off like some of the questions you have seen above, where you have to change and replace the code that looks like this:

```
# your code here
raise NotImplementedError # No Answer - remove if you provide an answer
```
In addition to this, with the autograded questions you can also run the tests to check if your answer is correct. You can run the tests as many times as you would like. If you see "success!" printed after running the autograding test cells that means your answer was correct. MDS labs will be a mix of manually graded and autograded questions.

Let's try our first autograded question below:

### Question 1.6.1
rubric={autograde:1}

In Python you can assign a value to an object using the equals sign, for example:

```
x = 5
```

creates an object named x that holds the value 5. 

In the cell below, create another object named y and assign it the value 10.

<!-- grade_id: cell-f4a194e73431293b, points: 0, grade: false, solution: true -->
```python
# create object named y with a value of 10
# YOUR CODE HERE
raise NotImplementedError()
print(y)
```

<!-- grade_id: cell-87a8d7a9f3fb5dad, points: 0, grade: true, solution: false -->
```python
# run this cell to check your answer in the cell above
assert sha1(str(y).encode('utf8')).hexdigest() == 'b1d5781111d84f7b3fe45a0852e59758cd7a87e5', "y should have the value of 10."
print("success!")
```

### Submit assignment

Run the cell below to submit your work:

```python
from canvasutils.submit import submit
submit(2039048, token_present=False, api_url="https://canvas.instructure.com/") # default url="https://canvas.ubc.ca/" which suits MDS but you can change this with argument api_url
```

### Attribution

- [Data 8: The Foundations of Data Science](http://data8.org/), specifically Lab 01
- [Software Carpentry](https://software-carpentry.org/), specifically the Unix Shell and Git lessons